<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2025/26
Prof. Dr. Heiner Giefers

# Praktikumstermin 2
1. Rückblick auf die Inhalte der letzten Kapitel
   - Schleifen und Funktionen
   - Dateien lesen und schreiben
   - Zeichenketten
2. Turtle Grafiken
   

### Schleifen und Funktionen

So gut wie jedes Programm verwendet Schleifen und Funktionen.
Es sind, neben der *bedingten Ausführung* (bzw. der *alternativen Ausführung*) die wichtigsten Elemente zur Strukturierung von Programmen.
In den kommenden Aufgaben, wollen wir noch einmal das Schreiben von Funktionen und Schleifen üben.

Funktionen in Python werden über das Schlüsselwort def definiert. Die Syntax einer Funktions-Definition sieht folgendermaßen aus:

```python
def myfunc(arg1, arg2,... argN):  
  '''Dokumentation'''  

  #Programmcode, optional mit return Anweisungen
```

Hier wird die Funktion `myfunc` definiert, welche mit den Parametern `arg1,arg2,....argN` aufgerufen wird.
Wir können bereits bei der Definition der Funktion Werte für einzelne Parameter angeben.
Jemand, der die Funktion verwendet, kann dann den Paramter weglassen und es wird der *default* Wert übernommen.
Die Parameter, die zwingend beim Aufruf der Funktion angegeben werden müssen, stehen vorne in der Liste, die mit *default*-Werten stehen hinten.

Wenn beim Programmaufruf Parameter angegeben werden, müssen diese in der gleichen Reihenfolge wie bei der Funktionsdefinition angegeben werden.
Sollen einige der *default*-Parameter übernommen werden, *fehlen* dieses Werte beim Funktionsaufruf und die Reihenfolge stimmt ab dem ersten übernommenen *default*-Parameter nicht mehr. In diesem Fall ist es notwendig, die folgenden Parameter, deren *default*-Werte überschrieben werden sollen, namentlich zu benennen.

Schauen wir uns die folgende Definition der Funktion `operator_tabelle` an.
Die Funktion besitzt drei Parameter:
`op` soll vom Typ `str` sein und die zu verwendende Operation angeben, `von` und `bis` legen den Start- bzw. Endwert der Operationstabelle fest.
Die Angabe eines Wertes für `op` ist von der Funktion gefordert. 
Für `von` und `bis` sind jeweils Werte vorgegeben, die aber nur verwendet werden, wenn entsprechende Argumente beim Funktionsaufruf weggelassen werden.

Was die Funktion tut und wie sie aufgerufen werden soll, wird auch im **`Docstring`** erklärt.
Dieser String kann am Anfang einer Funktion gesetzt werden.
Er dient zum Einen der besseren Verständlichkeit des Quellcodes.
Zum Anderen wird dieser String angezeigt, wenn die Python Hilfe-Funktion `help` mit dem Namen der Funktion aufgerufen wird.
Dies hilft den Personen, die die Funktion verwenden wollen, aber nicht unmittelbaren Zugriff auf den Quellcode haben.

**`return`** Anweisungen können grundsätzlich überall in der Funktion verwendet werden.
Es kann auch an mehreren Stellen ein `return`-Aufruf stehen, allerdings kann eine Funktion nur jeweils eine dieser Anweisungen ausführen.
Nach einem `return` wird die Funktion verlassen und bei einem erneuten Aufruf beginnt die Funktion wieder *von Vorne*.

```python
def operator_tabelle(op, von=1, bis=10):
    '''Die Funktion berechnet Additions- und Multiplikationstabellen.
    
    Parameters
    ----------
    op : str
            Der zu verwendende Operator, angegeben als String.
            Nur '+' und '*' werden unterstützt.

    von, bis: int, optional
            Start- und Endwert der Operationstabelle.
            
    Raises
    ------
    ValueError
            Wenn ein nicht unterstüzter Operator angegben wird.
    '''
    
    tabelle = ""
    if op=='+' or op=='*':
        for zeile in range(von, bis+1):
            for spalte in range(von, bis+1):
                if op=='+':
                    tabelle += f"{zeile+spalte:4d}"
                else:
                    tabelle += f"{zeile*spalte:4d}"
            tabelle += '\n'
        return tabelle
    else:
        raise ValueError("Fehler: Ich kann nur '+' oder '*' rechnen!")

s = operator_tabelle('*', bis=5)
print(s)
```

```python
help(operator_tabelle)
```

**Aufgabe 1:** Kann eine Python Funktion (je nach Aufruf) verschiedene Typen zurück liefern?

<!-- grade_id: cell-cda4d023f0634965, points: 0, grade: true, solution: true -->
```python
# Probieren Sie es einfach aus
# YOUR CODE HERE
raise NotImplementedError()
```

**Aufgabe 2:** Schreiben Sie eine Funktion `fak`, die eine ganze Zahl als Parameter entgegennimmt und die Fakultät dieser Zahl als Ergebnis zurückliefert. Für die Parameter soll weiterhin gelten:
- Falls der Parameter eine Fließkommazahl ist, soll diese zu einem Integer umgewandelt werden.
- Falls der Parameter weder vom Typ `int`, noch vom Typ `float`ist, soll die Funktion einen `ValueError` erzeugen.
- Wenn der Parameter weggelassen wird, soll standardmäßig die Fakultät von 42 berechnet werden.

<!-- grade_id: cell-3052fd6ef6e2e89b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ad5460e73674150e, points: 1, grade: true, solution: false -->
```python
### TEST ZELLE
assert fak(5)==120, "Bei der Fakultäts-Funktion stimmt etwas nicht"
assert fak()==1405006117752879898543142606244511569936384000000000, "Default Parameter verwenden"

try:
    etype=False
    fak('w')
except ValueError:
    etype = True
except Exception as e:
    pass
assert etype, "Bei falschem Paramter einen ValueError erzeugen"
```

**Aufgabe 3:** Schreiben Sie eine Funktion ggt(a,b), die den größten gemeinsamen Teiler (ggT) zweier Zahlen berechnet.


Hinweise: Eine effiziente Möglichkeit, den ggT zu berechnen ist der euklidische Algorithmus. Das Verfahren funktioniert folgendermaßen:

1. Solange `b!=0` (*b ungleich 0*) gilt, wiederhole folgende Schritte:
    - Berechne `r` als den ganzzahligen **Rest** der Division  `𝑎/𝑏`  (also: "r gleich a modulo b")
    - Setze `a=b` und `b=r`
1. Gebe a als Ergebnis zurück


<!-- grade_id: cell-bfeb0b8aa13f5cae, points: 0, grade: false, solution: true -->
```python
def ggt(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f3f15388f6852fe6, points: 1, grade: true, solution: false -->
```python
### TEST ZELLE
test_cases = [(123, 321, 3),(545,654, 109),(459,666, 9)]
assert all([ggt(a,b)==c for a, b, c in test_cases]), ''
```

**Zusatzaufgabe:**

Sie haben nun eine `ggt` Funktion, die den größten gemeinsamen Teiler zweier Zahlen berechnet.
Falls Sie den ggT von mehr als 2 Zahlen berechnen wollen, können Sie die Funktion `ggt` mehrfach benutzen.
Nehmen wir an Sie wollen den ggt der Zahlen `a`, `b`, und `c` berechnen.
Es gilt:

$ggT(a,b,c)=ggT(ggT(a,b),c)$.

Da für die ggT Funktion das Assoziativgesetz gilt, ist es unerheblich, in welcher "Reihenfolge" Sie die Funktion aufrufen.
Es gilt also:

$ggT(ggT(a,b),c)=ggT(a,ggT(b,c)$.

**Schreiben Sie eine Funktion `ggt_x`, die den ggT von 2 und mehr Zahlen berechnet.**

**Hinweise:** Verwenden Sie variable (Positions-) Parameter nach dem Muster `ggt_x(a, b, *weitere)`.

<!-- grade_id: cell-381f67e1bf0eb9bf, points: 0, grade: false, solution: true -->
```python
def ggt_x(a, b, *weitere):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f743fd5ad6da7f22, points: 0, grade: true, solution: false -->
```python
### TEST ZELLE
assert ggt_x(120, 540) == 60, 'Die Funktion soll auch mit 2 Parametern funktionieren.'
assert ggt_x(120, 540, 1200, 3420, 5620) == 20, 'Bei der Lösung stimmt etwas nicht.'
```

### Dateien lesen und Schreiben

Daten aus Dateien zu lesen oder zu schreiben benötigt man ständig. Wir wollen das anhand einer kleinen Aufgabe nochmal üben:

**Aufgabe 4a:** Schreiben Sie drei oder mehr Strings als Zeilen in eine Datei und schließen Sie die Datei. anschließend.

<!-- grade_id: cell-61d2494438025638, points: 0, grade: false, solution: true -->
```python
file_name = "ausgabe.txt"

fobj = None # Verwende fobj als Datei Handle
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-201ec32995367ce3, points: 0, grade: true, solution: false -->
```python
### TEST ZELLE
assert fobj.closed, 'Datei ist nicht geschlossen'
try:
    file = open(file_name, "r")
    assert file.read().count('\n')>2, '"ausgabe.txt" enthält weniger als 3 Zeilen'
    file.close()
except FileNotFoundError:
    print('"ausgabe.txt" ist nicht erstellt')
    raise
```

**Aufgabe 4b:** Öffnen Sie die Datei `ausgabe.txt` zum Lesen und geben Sie den Inhalt aus.

<!-- grade_id: cell-c37104af1f381b84, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-945b4331136cba10, points: 1, grade: true, solution: false -->
```python
### TEST ZELLE
from unittest.mock import patch
with patch('__main__.print') as mock_print:
     exec(_i)
mock_print.assert_called()
assert fobj.closed, 'File is not closed'
```

### Strings und Konvertierung

**Aufgabe 5:** Implementieren Sie eine Funktion `sternchen`, die als Argument einen String oder eine Zahl akzeptiert. Sie soll als Rückgabewert einen String liefern, der vorne und hinten je drei Sternchen und dazwischen, durch Leerzeichen abgetrennt, den Parameter enthält. Also z. B.:
```python
sternchen("Hallo") # soll "*** Hallo ***" zurückgeben
sternchen(42)      # soll "*** 42 ***"    zurückgeben
sternchen(4.123)   # soll "*** 4.123 ***" zurückgeben
```
Sie können für die Umwandlung von Zahlen in Strings die Funktion `str()` verwenden, sollen diese aber nur aufrufen, wenn der Parameter kein String ist. Bei der Implementierung wird Ihnen die Funktion `type()` helfen.

<!-- grade_id: cell-c1c7e574c742012d, points: 0, grade: false, solution: true -->
```python
def sternchen(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-21149defe9cc71ae, points: 1, grade: true, solution: false -->
```python
### TEST ZELLE
assert not [a for a in [15, 'Hallo'] if sternchen(a) != f'*** {a} ***'], 'Die Lösung ist nicht korrekt'
```

## Zeichnen mit `Turtle` (optional)


Logo ist eine Programmiersprache, die in den 1960-er Jahren am MIT entwickelt wurde und ursprünglich für den Einsatz in Schulen konzipiert war, um Kindern das Programmieren näher zu bringen.
Eine *Spezialität* von Logo sind sogenannte **Turtle-Grafiken**. Dabei handelt es sich um Strichzeichnungen, die entstehen, wenn sich eine *virtuelle Schiltkröte*, durch ein Programm gesteuert, über eine Zeichenfläche bewegt und dabei farbige Striche hinterlässt.
Auch für Python gibt es entsprechende Bibliotheken/Module, mit den Turtle-Grafiken erstellt werden können.

Das Ganze klingt sehr spielerisch, allerdings kann man mit Turtle Grafiken sehr gut Abläufe visualisieren, wie Sie mit Schleifen (für die Wiederholung von *Mustern*) und Funktionen (für das Kapseln von Mustern in wiederverwendbaren Blöcken) programmiert werden können.

`Turtle` kommt mit sehr wenigen Funktionen aus.
Allerdings gibt es, selbst für Python, verschiedenste Turtle Implementierungen, die jeweils einen unterschiedlichen Funktionsumfang bieten.

Eine Turtle-Version die sich in Jupyter-Notebooks anzeigen lässt ist `jupyturtle` und bietet nur sehr wenige Funktionen.
Hier eine Übersicht dazu:


- **`back(n)` oder `bk(n)`**: Bewegt die Schildkröte `n` Einheiten rückwärts.
- **`forward(n)` oder `fd(n)`**: Bewegt die Schildkröte `n` Einheiten vorwärts.
- **`left(n)` oder `lt(n)`**: Dreht die Schildkröte um `n` Grad nach links.
- **`right(n)` oder `rt(n)`**: Dreht die Schildkröte um `n` Grad nach rechts.
- **`make_turtle()`**: Erstellt eine neue Schildkröte und ein leeres Zeichenfeld.
- **`pen_up()`**: Hebt den Stift an, damit die Schildkröte keine Linie zeichnet.
- **`pen_down()`**: Senkt den Stift, damit die Schildkröte eine Linie zeichnet.
- **`toggle_pen()`**: Wechselt den Zustand des Stifts (hebt ihn, wenn er unten ist, und senkt ihn, wenn er oben ist).
- **`set_color(color)`**: Setzt die Farbe der gezeichneten Linien auf `color`.
- **`set_width(n)`**: Setzt die Breite der gezeichneten Linie auf `n`.
- **`jump_to(x, y)`**: Bewegt die Schildkröte ohne zu zeichnen zu den Koordinaten `(x, y)`.
- **`move_to(x, y)`**: Bewegt die Schildkröte zu den Koordinaten `(x, y)` und zeichnet dabei eine Linie.
- **`show()`**: Macht die Schildkröte sichtbar.
- **`hide()`**: Macht die Schildkröte unsichtbar.
- **`repeat(n, *items)`**: Wiederholt die angegebene Befehlsfolge `n` Mal.
- **`no_pen()`**: Hebt den Stift temporär an, während eines `with`-Blocks.
- **`draw()`**: Rendert die aktuelle Zeichnung auf dem Zeichenfeld.

In der folgenden Code Zelle werden die notwendigen Module importiert und die Zeichenfläche initialisiert.
Die Turtle-Bibliothek `jupyturtle` ist *objektorientiert*, das bedeutet, dass alle Operationen auf einem *Objekt der Klasse Turtle* ausgeführt werden müssen.
Um dies zu umgehen, werden in der folgenden Zelle die Turtle-Funktionen *um-definiert*, sodass Sie auf einem globalen Objekt ausgeführt werden.
Schauen Sie sich den Code gerne an, für die Lösung der Aufgaben können Sie ihn aber ignorieren.

```python
ON_JHUB = True
if ON_JHUB:
    try:
        from jupyturtle import *
    except:
        import sys
        !{sys.executable} -m pip install jupyturtle
        from jupyturtle import *
```

```python
from random import randint, choice, random, seed

rad = 300
XYSIZE = 2 * rad
STEPSIZE = 20

def rgb_to_hex(r, g, b):
    """Convert RGB values to HEX format."""
    return f'#{r:02x}{g:02x}{b:02x}'
        
def new_turtle(width=XYSIZE, height=XYSIZE):
    make_turtle(width=width, height=height)

def pencolor(t, r=-1, g=-1, b=-1): 
    if r == -1 and g == -1 and b == -1:
        return
    hex_color = rgb_to_hex(r, g, b)
    pen_color(hex_color)
```

**Aufgabe A:** Zeichnen Sie einen Fünfstern (*Pentagram*) in die Zeichenfläche. Einen Fünfstern können Sie *in einem Zug* zeichnen, indem Sie zuerst einen Schenkel zeichnen und dann in einen Winkel von 36° weiterzeichen. Diesen Winkel erreichen Sie, indem die Richtung der Schildkröte um 144° drehen (`180°-144°=36°`). Die Schenkellänge soll 100 Längeneinheiten betragen. Dach zeichnen Sie den nächsten Schenkel und drehen die Schildköte erneut um 144°. Diese Schritte wiederholen Sie insgesamt 5 Mal.

<!-- grade_id: cell-32a383d9cc2af58a, points: 0, grade: true, solution: true -->
```python
new_turtle()
# YOUR CODE HERE
raise NotImplementedError()
```

**Aufgabe B:** Schreiben Sie eine Funktion `star` zum Zeichnen eines Fünfsterns. Als Parmeter soll die Schenkellänge an die Funktion Übergeben werden können. Fehlt der Parameter beim Aufruf, soll die Schenkellänge zufällig im Bereich 10 bis 50 Längeneinheiten gewählt werden.
Die Ecken der Sterne sollen in zufällige Richtungen (und nicht immer in die gleiche Richtung) zeigen. Dazu drehen Sie vor dem Zeichnen die Schildkröte um einen zufälligen Winkel.

Um eine gleichverteilte Zufallszahl im Bereich `[a,b]` zu erhalten, verwenden Sie die Funktion `randint(a,b)`.

<!-- grade_id: cell-ae4766290d21eda6, points: 0, grade: true, solution: true -->
```python
new_turtle()
# YOUR CODE HERE
raise NotImplementedError()

star(100)
star()
star()
```

**Aufgabe C:** Schreiben Sie eine Funktion `draw_a_star` zum Zeichnen eines Fünfsterns an eine zufällige Position auf der Zeichenfläche. Die Farbe des Sterns soll ebenfalls zufällig gewählt werden.

Für die Position des Sterns kommen die `x`- und `y`-Koordinaten zwischen `-rad` und `rad` infrage. Die Farbe für die Strichkanten wird nach dem RGB-Modell gebildet. Dazu müssen die Rot-/Grün-/Blau-Anteile jeweils in Prozent (im Bereich `[0,1]`) angegeben werden. Eine Fließkommazahl in diesem Bereich erhalten Sie mit dem Aufruf ` random()`.

<!-- grade_id: cell-c87f0f36d5a009a8, points: 0, grade: true, solution: true -->
```python
new_turtle()
# YOUR CODE HERE
raise NotImplementedError()

for i in range(10):
    draw_a_star()
```